Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

---

### HW 4: Naive Bayes Classification


# Part 1 - Theory / Written Problems


### [16 points] Problem 1 - Naive Bayes for Tennis Prediction 
***

Suppose you are trying to learn a person's decision whether to play tennis or not on a given day using features corresponding to precipitation forecast, temperature, humidity, and wind. You're given the following training data: 

$$
\begin{array}{c|c|c|c|c}
\textbf{Forecast} & \textbf{Temp} & \textbf{Humidity} & \textbf{Wind} & \textbf{PlayTennis} \\
\hline 
\textrm{sunny} & \textrm{hot} & \textrm{high} & \textrm{weak} & \textbf{No} \\ 
\textrm{sunny} & \textrm{hot} & \textrm{high} & \textrm{strong} & \textbf{No} \\ 
\textrm{overcast} & \textrm{hot} & \textrm{high} & \textrm{weak} & \textbf{Yes} \\ 
\textrm{rainy} & \textrm{mild} & \textrm{high} & \textrm{weak} & \textbf{Yes} \\ 
\textrm{rainy} & \textrm{cool} & \textrm{normal} & \textrm{weak} & \textbf{Yes} \\ 
\textrm{rainy} & \textrm{cool} & \textrm{normal} & \textrm{strong} & \textbf{No} \\ 
\textrm{overcast} & \textrm{cool} & \textrm{normal} & \textrm{strong} & \textbf{Yes} \\ 
\textrm{sunny} & \textrm{mild} & \textrm{high} & \textrm{weak} & \textbf{No} \\ 
\textrm{sunny} & \textrm{cool} & \textrm{normal} & \textrm{weak} & \textbf{Yes} \\ 
\textrm{rainy} & \textrm{mild} & \textrm{normal} & \textrm{weak} & \textbf{Yes} \\ 
\textrm{sunny} & \textrm{mild} & \textrm{normal} & \textrm{strong} & \textbf{Yes} \\ 
\textrm{overcast} & \textrm{mild} & \textrm{high} & \textrm{strong} & \textbf{Yes} \\ 
\textrm{overcast} & \textrm{hot} & \textrm{normal} & \textrm{weak} & \textbf{Yes} \\ 
\textrm{rainy} & \textrm{mild} & \textrm{high} & \textrm{strong} & \textbf{No} \\ 
\end{array}
$$

**Part A (2pt)**: Estimate the **unconditional** probabilities $p(\textrm{PlayTennis=Yes})$ and $p(\textrm{PlayTennis=No})$ from the training data. These are also called the "prior" probabilities. **Note in all the problems that follow, unless directed otherwise, you must show/describe how you arrived at your answer, even if the method is simple. Simply writing  the numeric answer is not enough.**



<!-- grade_id: cell-e1f2324117ff87fd, points: 2, grade: true, solution: true -->
$p(\textrm{PlayTennis=Yes}) = \frac{9}{14}$ and $p(\textrm{PlayTennis=No}) = \frac{5}{14}$

**Part B (6 pt)**: 
1. For each feature, estimate the class-conditional probabilities $p(\textrm{feature value} \mid \textrm{PlayTennis} = "Yes" )$ and $p(\textrm{feature value} \mid \textrm{PlayTennis} = "No" )$. Show your work.

2. For a given probability mass function for a discrete random variable $X$, the sum of the probabilities over the range of values the variable can take must add up to 1. Is the following equation true? Why or why not?

$$  p(\textrm{Forecast=Sunny} | \textrm{Humidity = high} ) +  p(\textrm{Forecast=Overcast} | \textrm{Humidity = low} ) +  p(\textrm{Forecast=rainy} | \textrm{Humidity = normal} ) = 1 $$



<!-- grade_id: cell-2584fb81679aaf31, points: 6, grade: true, solution: true -->
2. NO, this is not the case,the probability would be one if we summed up 
$p(Forecast=Sunny | Humidity=high) + p(Forecast=Overcast | Humidity = high) + p(Forecast=rainy | Humidity = high)$

**Part C (4 pt)**:
    
1. Consider the Naive Bayes model where the $x_i$ features are discrete random variables (also known as Multinomial Naive Bayes): 

$$ p(y \mid x_1, ..., x_n) \approx p(y) \Pi_i^n p(x_i | y) $$

In this case, our prediction is made by:

$$ \hat{y}(x) = \mathop{\mathrm{argmax}}_y p(y) \Pi_i^n p(x_i \mid y) $$

What will happen if we receive a query point $x$ such that the $i$-th feature $x_i$ contains a value not in the original dataset. Suppose this value is $v$. What is $p(x_i=v \mid y)$? Such a case may occur if, for example, our features are words in a sentance. Then a rare word may appear that was not in the original training dataset. Assume you trained your model only on words in the training data using exact counts of features to estimate the requisuite probabilities. 

<!-- grade_id: cell-1c0e5ee75b537eef, points: 4, grade: true, solution: true -->
Then we must estimate the label based off of the posterior! 

$p(z | x_{i} ; \theta) = \frac{p(x_{i} | z ; \theta)p(z ; \theta)}{\Sigma_{z^{'} = 1}^{k} p(x_{i} | z^{'}); \theta)p(z^{'} ; \theta) }$

**Part D (4 pts)**: 
One way to solve the problem you discovered in (1) is to artificially inflate the count of all possible words by some small number. So in an NLP dataset, all possible words would have at least some minimum count when calculating the above model. We can also introduce an `<unk>` word with a minimum count that is assigned to any words that we didn't account for at train-time. Inflating the counts results in so-called **Laplace Smoothing**:

$$p(x_i =j| y_i=y) \approx \frac{N_{yij}+ \alpha }{N_y + \alpha D} $$

where $N_{y}$ is the total count of all samples with label $y$, $N_{yij}$ is the total count of all samples such that $x_i=j$ and $y_i=y$, and $D$ is the number of features. $\alpha$ is an adjustable parameter and can be thought of as a an additional factor added to each feature count, although it need not be an integer.

Assume a Naive Bayes model with Laplace Smoothing such that $\alpha=2$. What would your Naive Bayes model predict for the following two weather conditions? Show all work.  
- **Forecast**=overcast, **Temp**=cool, **Humidity**=high, **Wind**=weak  
- **Forecast**=sunny, **Temp**=hot, **Humidity**=normal, **Wind**=strong  

<!-- grade_id: cell-762f9fd576f071b9, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

### PART 2 [15 points] Implementing Naive Bayes for Text Classification 
***

In this problem you'll implement a general Discrete Naive Bayes class for text classification. Your tasks will be to implement `train`, `predict_log_score`, and `predict` routines to learn the Naive Bayes model parameters from a collection on text and make predictions on unseen validation data. 

The skeleton for the `TextNB` class is below. 

**Important Note**: In Problem 3 we'll be using the `TextNB` class to make predictions about Twitter data.  Since real-world text data typically has a large number of features, you'll want to make your implementation reasonably efficient so that your experiments in Part 3 don't take forever. 


```python
import numpy as np
import ast
import matplotlib.pyplot as plt
from sklearn.feature_extraction.text import CountVectorizer
```

<!-- grade_id: cell-81e29b259b0173d3, points: 0, grade: false, solution: true -->
```python
class TextNB:
    def __init__(self, text_train, y_train, alpha=1.0):
        """
        :param text_train: a list or ndarray of text strings to use as training data 
        :param y_train: an ndarray of true labels associated with the text data 
        :param alpha: the Laplace smoothing parameter 
        """
        
        # store training data 
        self.text_train = text_train 
        self.y_train = y_train 
        
        # store smoothing parameter
        self.alpha = alpha 
        
        # get number of classes 
        self.num_classes = len(np.unique(y_train))
        
        # initialize vocab to feature map 
        self.vocab = dict() 
        
        # initialize class counts 
        self.class_counts = np.zeros(self.num_classes, dtype=int)
        
        # initialize feature counts (Note, will need to update this with the correct
        # number of columns during the training process)
        self.feature_counts = np.zeros((self.num_classes, 0), dtype=int)
        
        # Any extra fields you want go here
        # YOUR CODE HERE
        self.word_counter = {}
        
        self.V = CountVectorizer()
        
    def train(self):
        """
        Learn the vocabularly, class_counts, and feature counts from the training data 
        """
        
        # YOUR CODE HERE
        for word in self.text_train:
            if word in self.word_counter:
                self.word_counter[word] += 1
            else:
                self.word_counter[word] = 1
        
        X = self.V.fit(self.text_train)
        for i, row in enumerate(X):
            self.vocab[self.text_train[i]] = row
            
                    
                    
    def predict_log_score(self, text_str):
        """
        Get the log-probability score for each class
        for a query string
        
        :param text_str: a single string of text to compute the log_score for 
        """
        
        # YOUR CODE HERE
        raise NotImplementedError()
        
    
    def predict(self, text_list):
        """
        Predict the class of each example in text_list  
        
        :param text_list: a list or ndarray of text strings to make predictions on 
        """
        
        # YOUR CODE HERE
        raise NotImplementedError()

```

**Not that in all of the following you may NOT use scikit-learn NaiveBayes classes.**

**Part A**: Complete the `train` function in the `TextNB` class to prepare to make Naive Bayes predictions using Laplace smoothing.  In this routine you will need to populate the following data structures: 

`self.vocab`: A Python dictionary that maps distinct terms found in the training set to unique indices in $\{0, 1, \ldots, |V|-1\}$.  This will allow us to quickly look up frequency counts for an encountered term in a Numpy array. Note that while the data is fairly clean (We've removed punctuation and made all characters lowercase) you should take care that you're not accidentally counting whitespace in the vocabulary. 

`self.class_counts`: A 1D Numpy array of length `self.num_classes` which counts the number of documents in the training set that belong to each class. 

`self.feature_counts`: A 2D Numpy array of dimensions `self.num_classes` $\times$ $|V|$. The $(c,k)$-entry in this array should be the number of times that term $k$ appears in documents belonging to class $c$. Note that we're using the Bag-of-Words approach here, so if a term appears multiple times in a single document, each instance of that term should be counted.  

When you think you're done, execute the following unit test. 2 are available for you, one is hidden.

<!-- grade_id: cell-e925aff87ac8cf0b, points: 4, grade: true, solution: false -->
```python
def testVocab(): 
        """
        test vocabulary 
        """
        text_train = np.array(["work buy  money", "prince opportunity viagra", "fly   buy prince", "money buy fly", "fly home prince"])
        y_train = np.array([0, 1, 1, 1, 0], dtype=int)
        nb = TextNB(text_train, y_train)
        nb.train()
        vocab_set = set(["work", "buy", "money", "prince", "opportunity", "viagra", "fly", "home"])
        for w in vocab_set:
            assert(w in nb.vocab)
            assert(nb.vocab[w] in set(list(range(8))))
            
def testClassCounts(): 
        """
        test counts of documents in each class 
        """
        text_train = np.array(["work buy  money", "prince opportunity viagra", "fly   buy prince", "money buy fly", "fly home prince"])
        y_train = np.array([0, 1, 1, 1, 0], dtype=int)
        nb = TextNB(text_train, y_train)
        nb.train()
        vocab_set = set(["work", "buy", "money", "prince", "opportunity", "viagra", "fly", "home"])
        assert(nb.class_counts[0] == 2)
        assert(nb.class_counts[1] == 3)

testVocab()
testClassCounts()
```

<!-- grade_id: cell-acbeb3b7ea9f7c20, points: 2, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

**Part B**: Complete the `predict_log_score` function in the `TextNB` class to take in a single string of text and compute the associated log-score for each class. For now, you should use Laplace smoothing with $\alpha=1$ for both the class-priors and the class-conditional probabilities.  In **Problem 3** we'll experiment with different variants of Laplace smoothing, so if you like you can read ahead now and implement the general version of Laplace smoothing from the beginning.  

**Note**: For simplicity and testing purposes, do not implement an `UNK` feature.  Instead, if you encounter a term not in the vocabulary you can safely ignore it. 

When you think your `predict_log_score` function is working well, execute the following unit test. Some unit tests are hidden. 

<!-- grade_id: cell-58f78a77269133ab, points: 4, grade: true, solution: false -->
```python
def testPrediction(): 
    """
    test vocabulary 
    """
    text_train = np.array(["work buy  money", "prince opportunity viagra", "fly   buy prince", "money buy fly", "fly home prince"])
    y_train = np.array([0, 1, 1, 1, 0], dtype=int)
    nb = TextNB(text_train, y_train, alpha=1.0)
    nb.train()
    test1 = "money money money"
    test2 = "money fly buy"
    logits1 = nb.predict_log_score(test1)
    logits2 = nb.predict_log_score(test2)
    print(logits1, logits2)
    diff = np.abs(logits1-np.array([-6.685, -6.9798]))
    assert(diff.sum()<1e-4)
    diff = np.abs(logits2-np.array([-6.685, -6.1689]))
    assert(diff.sum()<1e-4)
    
testPrediction()
```

**Part C**: Finally, implement the `predict` method to take in a list or ndarray of text data, call `predict_log_score`, and return a vector of predicted labels. 

When you think you're done, execute the following unit test. 

<!-- grade_id: cell-436849047d2409a9, points: 5, grade: true, solution: false -->
```python
def testPredict(): 
    text_train = np.array(["work buy  money", "prince opportunity viagra", "fly   buy prince", "money buy fly", "fly home prince"])
    y_train = np.array([0, 1, 1, 1, 0], dtype=int)
    nb = TextNB(text_train, y_train)
    nb.train() 
    yhat = nb.predict(text_train)
    assert(len(yhat)== 5)
    assert(all(yh==yi for yh, yi in zip(yhat, y_train)))
        
testPredict()
```

### [12 points] Problem 3: Predicting the Sentiment of Tweets sent from Passengers to Airlines 
***

In this problem you'll use the `TextNB` class you wrote in **Problem 2** to make predictions about the sentiment of tweets sent by passengers to airlines.  Execute the following cell to load the data. 

```python
def get_dataset(path='airline_tweets.pklz'):
    f = open('airline_tweets.pklz','rb')
    text_train, y_train, text_valid, y_valid, text_all, y_all = np.load(f, allow_pickle=True)
    f.close()
    return {
        'train':{
            'tweets': text_train,
            'labels': y_train
        },
        'valid': {
            'tweets': text_valid,
            'labels': y_valid
        },
        'all': {
            'tweets': text_all,
            'labels': y_all
        }
    }

# This should be the only global variable you use in your cells.
# Treat it as immutable. Do not access it directly in any function, it only gets
# passed in.
global_dataset = get_dataset()
```

**Part A (4Pts)**: Explore the data and fill in the below function toanswer the following questions: 

- How many total examples are there in the training set?
- How many total examples are there in the validation set?
- Which binary label ($\{0,1\}$) corresponds to tweets with positive sentiment?
- What percentage of tweets in the training set have true positive sentiment? 
- What percentage of tweets in the validation set have true positive sentiment? 

```python
### This cell is for scratch space to explore the data before filling in the below function
```

<!-- grade_id: cell-1a9a6c6a18b5e01b, points: 0, grade: false, solution: true -->
```python
def get_dataset_info(dataset):
    '''
    dataset - the loaded dataset
    
    returns:
        np.array of size 5 with the elements from above (in order)
    '''
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-44d6f1137fdafd2f, points: 4, grade: true, solution: false -->
```python
# DO NOT USE OR DELETE THIS CELL

```

**Part B (4 Pts)**: Use your `TextNB` class to learn a Naive Bayes classifier for the airline Twitter data.  What accuracy do you achieve on the training set and what accuracy, precision, and recall do you achieve on the validation set? 

You may use `sklearn.metrics`.

<!-- grade_id: cell-f5c5782edb8ff5d8, points: 0, grade: false, solution: true -->
```python
import sklearn.metrics as metrics

def build_and_evaluate(dataset, alpha=1.0):
    '''
    args - 
        dataset: the dataset object
        alpha: the alpha value (for Laplace smoothing) given to TextNB
    
    Returns an np.array of 6 numbers: 
        train set accuracy,
        train set precision, 
        train set recall, 
        valid set accuracy, 
        valid set precision,
        valid set recall
        
    You should use your TextNB class
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bb7c76ec22048a6f, points: 4, grade: true, solution: false -->
```python
# DO NOT USE OR DELETE THIS CELL

```

**Part D**: Write a procedure to perform $K$-Folds cross-validation on the entire data set (`train_all` and `y_all`) to estimate the accuracy of your NB classifier for various values of $\alpha$ and make a plot showing your results.  

To do the partitioning into folds we recommend leveraging sklearn's [StratifiedKFold](https://scikit-learn.org/stable/modules/cross_validation.html#cross-validation-iterators-with-stratification-based-on-class-labels) routine.  The documentation demonstrates how it can be used.  

For your plot, use at least $K=5$ folds and at least $5$ different values of $\alpha$ between $0.1$ and $1.5$ (or higher).  Which value of $\alpha$ seems to perform the best? 

<!-- grade_id: cell-7ef6a941fdcb5f81, points: 4, grade: true, solution: true -->
```python
from sklearn.model_selection import StratifiedKFold

def kfold_validation(dataset, alpha_vals=np.linspace(0.1, 1.5, 5)):
    '''
    Use your build_and_evaluate function to perform kfold_validation
    to select the alpha parameter.
    
    You should return a A x 6 array, 1 row for each alpha value,
    and 1 column for each metric returned from build_and_evaluate
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1bbf1c7c45a697ea, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-69661fdbdfd6efcb, points: 1, grade: true, solution: false -->
```python
# Do not modify this cell. 
```

```python

```